## 1) What, in your opinion, was the most significant advancement in the history of computers? What would a next big leap be?

<!-- grade_id: cell-2b7b7a3717ed7c28, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## 2) Define: algorithm; high level language; bit.

<!-- grade_id: cell-b0950da5a192e035, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## 3) Programming: Write a program that lists all the multiples of 15 greater than −200 and less than 2000.

<!-- grade_id: cell-640c6f13e30f6fbb, points: 0, grade: false, solution: true -->
```python
# define an empty list object called list then append the list with multiples of 15. Be careful where you 
# start and where you stop.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-70172c44c90ecb80, points: 3, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(list,[-195, -180, -165, -150, -135, -120, -105, -90, -75, -60, -45, -30, -15, 0, 15, 30, 45, 60, 75, 90, 105, 120, 135, 150, 165, 180, 195, 210, 225, 240, 255, 270, 285, 300, 315, 330, 345, 360, 375, 390, 405, 420, 435, 450, 465, 480, 495, 510, 525, 540, 555, 570, 585, 600, 615, 630, 645, 660, 675, 690, 705, 720, 735, 750, 765, 780, 795, 810, 825, 840, 855, 870, 885, 900, 915, 930, 945, 960, 975, 990, 1005, 1020, 1035, 1050, 1065, 1080, 1095, 1110, 1125, 1140, 1155, 1170, 1185, 1200, 1215, 1230, 1245, 1260, 1275, 1290, 1305, 1320, 1335, 1350, 1365, 1380, 1395, 1410, 1425, 1440, 1455, 1470, 1485, 1500, 1515, 1530, 1545, 1560, 1575, 1590, 1605, 1620, 1635, 1650, 1665, 1680, 1695, 1710, 1725, 1740, 1755, 1770, 1785, 1800, 1815, 1830, 1845, 1860, 1875, 1890, 1905, 1920, 1935, 1950, 1965, 1980, 1995])
```

## 􏰀4) Programming: Write a program that calculates the sum
$$\sum_{k=-100}^{75}\frac{k^3-2k}{k^2+2}$$
## Can this sum be computed analytically?

<!-- grade_id: cell-294356bd9b3e974b, points: 0, grade: false, solution: true -->
```python
# define a function that contains your formula called f, then define a sum=0, then define
# a loop that adds f(k) onto your sum. Check the range of your loop
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-87d41668f23b9a81, points: 3, grade: true, solution: false -->
```python
assert_equal(sum,-2198.856218040634)
```

## 5) Programming: Write a program that determines whether the sequence 
$$a_n=\frac{(\cos(n))^2}{2^n}$$
## converges or diverges. Make a plot and describe the behavior.

<!-- grade_id: cell-17b78fde82a13eb6, points: 0, grade: false, solution: true -->
```python
# Define a function called a that contains your function formula. You need numpy library
# in order to use the cosine function numpy.cos()
# initiate an empty list object called list. Append it with a(n) for n=1 to 50
# finally plot your result
import numpy
import matplotlib.pyplot as plt
# YOUR CODE HERE
raise NotImplementedError()
# once you are done, print a description of the type of behavior of your sequence, 
# for example
print("This sequence converges to zero, or this sequence diverges to infinity, or there\
    are two subsequential limits, etc.")
```

<!-- grade_id: cell-e21ace443233c287, points: 3, grade: true, solution: false -->
```python
assert_equal(a(50),8.270357164830974e-16)
```